// intro1.rs
// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise, even
// after you already figured it out. If you got everything working and feel
// ready for the next exercise, remove the `I AM NOT DONE` comment below.
// Execute `rustlings hint intro1` or use the `hint` watch subcommand for a hint.
//
// If you're running this using `rustlings watch`: The exercise file will be reloaded
// when you change one of the lines below! Try adding a `println!` line, or try changing
// what it outputs in your terminal. Try removing a semicolon and see what happens!

fn main() {
    println!(r#"                                ,,     @m     ╓                                 "#);
    println!(r#"                          ╖    ╒╠╠▒, ╒╣╠╠▒  φ╠╠▒    ╓┐                          "#);
    println!(r#"                         @╠╠▒╗g╣╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒╓@╠╠╠                          "#);
    println!(r#"                   ]╣▒╗,,╬╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒,g@▒▒                    "#);
    println!(r#"                   ]╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠                    "#);
    println!(r#"              %▒▒@╗@╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╗@▒▒▒               "#);
    println!(r#"               ╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒               "#);
    println!(r#"          ┌╓╓╓╓╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒--╓╓           "#);
    println!(r#"          ╙╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╬           "#);
    println!(r#"           ╚╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠            "#);
    println!(r#"       φ@▒▒╬╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒▒@╗,       "#);
    println!(r#"       └╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠██▄  ╙╠╠╠╠╠╠╠╟██  "╬╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╩        "#);
    println!(r#"        ,╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╝▐██▌   └╠╠╠╠╠⌐╙██▌   ╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒         "#);
    println!(r#"       @╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒ '└     ╠╠╠╠╠        ╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒m       "#);
    println!(r#"    .@╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠m      ╠╠╠╠╠╠▒     .@╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒m    "#);
    println!(r#"   @╠╠╠╠╝└░░╙╚╚╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠▒▒▒▒╠╠╠╠╠╠╠╠╠▒▒▒╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠╠░=░░╚╙╠╠╠╠╠   "#);
    println!(r#"   ^╠╠╠╠╠  "░≥ `╙╠╠╠╠╚╩╠╠╠║╙╓╔φφg╖└╙║╠╠╠╠╠╠╠╙╓φ▒▒▒@╗╙╝╠╠╠╩╚╙╠╠╠╠╝   ░╙ .╠╠╠╩    "#);
    println!(r#"     `╚╠╠╠,  7φ   ╚╠╠╠░ⁿ=╓@╠╠╠╠╠╠╠╠▒σ└╚╩╩╝░φ╠╠╠╠╠╠╠╠╠▒Ç="7φ╠╠╠╩"   ░   ╠╠╠┘     "#);
    println!(r#"        ╚╠╠m   φ    ╚╠╠▒▒╠╠╠╠╠╠╠╠╝╝╝╚╚≈   └"""└╙╙╚╠╠╠╠╠▒╠╠╠╩╙    ,░   @╠╩       "#);
    println!(r#"          ╙╠▒   `     `╙╚╠╠╠╠╠ε                  ╔╠░░░░╠╙            @╠"        "#);
    println!(r#"            └║ε           ╚░░░░░╠φε,         .╔δ░░░░░╚"             @╜          "#);
    println!(r#"              "ε            ⁿ╠░░░░░░╚`    ╙╚░░░░░╚╙"               ╒⌐           "#);
    println!(r#"                                ""`                                             "#);
    println!(r#" _       __     __                             __           ____             __ "#);
    println!(r#"| |     / /__  / /________  ____ ___  ___     / /_____     / __ \__  _______/ /_"#);
    println!(r#"| | /| / / _ \/ / ___/ __ \/ __ `__ \/ _ \   / __/ __ \   / /_/ / / / / ___/ __/"#);
    println!(r#"| |/ |/ /  __/ / /__/ /_/ / / / / / /  __/  / /_/ /_/ /  / _, _/ /_/ (__  ) /_  "#);
    println!(r#"|__/|__/\___/_/\___/\____/_/ /_/ /_/\___/   \__/\____/  /_/ |_|\__,_/____/\__/  "#);
    println!();
    println!("This exercise compiles successfully. The remaining exercises contain a compiler");
    println!("or logic error. The central concept behind Rustlings is to fix these errors and");
    println!("solve the exercises. Good luck!");
    println!();
    println!("The source for this exercise is in `exercises/intro/intro1.rs`. Have a look!");
    println!("Going forward, the source of the exercises will always be in the success/failure output.");
}
